---
title: Теоретическая глава
description: This is a template chapter.
---

## Будет ли сегодня дождь?

```yaml
type: PureMultipleChoiceExercise
key: 4ff2227cf9
xp: 50
skills: 1
```
Известно, что каждый день дождь может либо идти, либо не идти. Каждое утро исследователь Геннадий включает специальный прибор, который выводит "1", если дождь в этот день будет, и "0", если дождя не будет. Выберите верное утверждение о работе прибора.

`@possible_answers`

- [Значение прибора до его включения – случайная величина.]
- Значение прибора до его включения – константа.
- Геннадий увидел, что прибор показывает "1". Значение прибора – случайная величина. 
- В прошлый четверг дождя не было, и утром прибор показал "0". Показанное значение прибора – случайная величина. 

`@hint`
Случаная величина может принимать некоторые значения с определенной вероятностью. Реализация случайной величины в данный момент времени – константа. 

`@feedback`
- Верно! До того, как прибор включён, Геннадий может только предположить, какое значение он покажет.
- Геннадий не знает, какую константу покажет прибор до включения.
- Геннадий знает, какое значение показал прибор, случайность отсутствует. 
- Геннадий знает, какое значение показал прибор, случайность отсутствует.




---
## В ожидании непогоды

```yaml
type: NormalExercise
key: 06813252d3
lang: r
xp: 100
skills: 1
```


`@instructions`

Как вы поняли, значение прибора до включения – случайная величина, имеющая распределение Бернулли:
```{r fig.width=1, fig.height=10,echo=FALSE}
library(png)
library(grid)
img <- readPNG("https://dl.dropboxusercontent.com/s/ocxkserxtg21gax/1_1.png?dl=0")
grid.raster(img)
```


sdf
fgj

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
