---
title       : Работа с датасетом
description : Insert the chapter description here


---
## Знакомство с датасетом

```yaml
type: NormalExercise
key: e21d6ca1b9
lang: r
xp: 100
skills: 1
```
Перед работой с любым датасетом его необходимо обработать и ознакомиться с ним. Датасет `candy` уже загружен в Ваше рабочее пространство.

Данный датасет относительно небольшой, но в будущем Вы можете столкнуться с датасетами, в которых будет более тысячи наблюдений и десятки переменных, поэтому для примерного знакомства Вам будет достаточно первых несколько наблюдений. Встроенная функция 'head()' выводит в консоль первые 6 наблюдений, что позволяет ознакомиться со всеми переменными и сложить представление о датасете, однако количество выводимых наблюдений можно уточнить внутри самой функции. 

Иногда взгляда на первые наблюдения может быть недостаточно, тогда для анализа датасета можно использовать функцию 'describe()' из библиотеки psych, которая формирует краткий отчёт по каждой переменной. Вы можете подробнее ознакомиться с её выводом, выполнив задание.

`@instructions`

- Загрузите библиотеку psych с помощью функции library()
- Проанализируйте данный датасет.
- Выгрузите первые 10 наблюдений.

`@hint`
- Используйте функции describe() и head()

`@pre_exercise_code`
```{r}
candy <- read.csv(url("https://raw.githubusercontent.com/fivethirtyeight/data/master/candy-power-ranking/candy-data.csv"))
```

`@sample_code`
```{r}
# подключите библиотеку psych
___(psych)

# проанализируйте данный датасет
___(___)

# выгрузите первые 10 наблюдений
___(___, 10)
```

`@solution`
```{r}
# подключите библиотеку psych
library(psych)

# проанализируйте данный датасет
describe(candy) 

# выгрузите первые 10 наблюдений
head(candy, 10)
```

`@sct`
```{r}
#first instruction
test_student_typed("library(psych)", not_typed_msg = "Вы не подгрузили библиотеку psych. Внимательно прочтите инструкцию")

#second instruction
test_student_typed("describe(candy)", not_typed_msg = "Возникла проблема с анализом датасета. Внимательно прочтите инструкцию.")

#third instruction
test_student_typed("head(candy, 10)", not_typed_msg = "Возникла проблема с выгрузкой первых 10 наблюдений. Внимательно прочтите инструкцию.")

#General
test_error()
success_msg("Отлично! Пришло время проверить, как хорошо Вы познакомились с датасетом.")
```



---
## Анализ датасета

```yaml
type: MultipleChoiceExercise
key: 6a08fb4f1d
lang: r
xp: 50
skills: 1
```
Какое из этих утверждений верное?

`@instructions`

- В этом датасете 79 наблюдений
- Самая высокая оценка шоколадки - 91%
- В среднем в этих шоколадках 47% сахара
- В этом датасете 12 переменных

`@hint`

- Используйте упомянутую ранее функцию describe()

`@pre_exercise_code`
```{r}
candy <- read.csv(url("https://raw.githubusercontent.com/fivethirtyeight/data/master/candy-power-ranking/candy-data.csv"))
library(psych)

```

`@sct`
```{r}
msg1 <- "Неверно. Посмотрите на значение n."
msg2 <- "Не совсем. Посмотрите на максимальное значение переменной winpercent."
msg4 <- "Неправильно. Посмотрите, сколько названий переменных Вы видите при анализе датасета."
msg3 <- "Замечательно! Теперь Вы познакомились с датасетом и готовы начать более подробный анализ."

test_mc(3, feedback_msgs = c(msg1, msg2, msg3, msg4))

#General
test_error()
success_msg("Замечательно! Теперь Вы познакомились с датасетом и готовы начать более подробный анализ.")
```


---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: e7fac09106
lang: r
xp: 100
skills: 1
```
Из предыдущих глав Вы можете помнить, что такое выборочные характеристики, пришло время отработать эти знания на датасете. 

Следует помнить, что некоторые выборочные характеристики (например, среднее) можно считать только по числовым переменным. Если же в Вашем датасете _df_ переменная _gender_ принимает значения "М" и "Ж", то Вы не сможете посчитать долю женщин, но если эта переменная будет принимать значения "0" и "1", то Вы сможете воспользоваться функцией  `mean(df$gender)`.

`@instructions`

- Посчитайте среднюю долю карамели в конфетах.
- Посчитайте дисперсию доли шоколада в конфетах.
- Посчитайте медиану процента стоимости конфет.


`@hint`

- Если Вы забыли названия переменных, загрузите библиотеку psych и воспользуйтесь функцией describe().

`@pre_exercise_code`
```{r}
#candy <- read.csv(url("https://raw.githubusercontent.com/fivethirtyeight/data/master/candy-power-ranking/candy-data.csv"))
```

`@sample_code`
```{r}
# Средняя оценка карамельности
___(___$___)

# Дисперсия оценки шоколадности
___(___$___)

# Медиана стоимости конфет
___(___$___)
```

`@solution`
```{r}
# Средняя оценка карамельности
#mean(candy$caramel)

# Дисперсия оценки шоколадности
#sd(candy$chocolate)

# Медиана стоимости конфет
#median(candy$pricepercent)
```

`@sct`
```{r}
#first instruction
#test_student_typed("mean(candy$caramel)", not_typed_msg = "Возникла проблема с оценкой карамельности. Вы воспользовались функцией mean() для нужной переменной?")

#second instruction
#test_student_typed("sd(candy$chocolate)", not_typed_msg = "Возникла проблема с оценкой шоколадности. Вы воспользовались функцией sd() для нужной переменной?")

#third instruction
#test_student_typed("median(candy$pricepercent)", not_typed_msg = "Возникла проблема с оценкой медианы стоимости конфет. Вы воспользовались функцией median() для нужной переменной?")

#General
#test_error()
#uccess_msg("Великолепно! Однако подобное преставление не очень наглядно, в следующих упражнениях у Вас будет возможность визуализировать некоторые показатели.")
```


