---
  title: "Выборочные характеристики"
  description: "This is a template chapter."
  v2: true

---


```yaml
type: undefined

xp: 0

key: a85baa952c



```














---
## Интерпретация доверительных интервалов

```yaml
type: PureMultipleChoiceExercise

xp: 50
skills: 1
key: e692d3bd0c



```

Выберите верное суждение о двустороннем 99% доверительном интервале.







`@possible_answers`
- Построив 99 % доверительный интервал, мы получаем интервал, в котором лежит 99 % наших наблюдений.
- С вероятностью 99% неизвестный параметр  принимает значения из полученного доверительного интервала.
- [Проведя эксперимент большое число раз, мы ожидаем, что в 1% экспериментов доверительный интервал не будет накрывать истинное значение параметра.]
- Вероятность, что неизвестный параметр примет значение больше верхней границы доверительного интервала, равна 1%.

`@feedbacks`
- В общем случае, для каждой конкретной выборки в границы доверительного интервала может попасть различное число наблюдений. Вам важно оценить, как часто в эти границы попадает неизвестный параметр.
- Неизвестный параметр, который вы пытаетесь оценить, принимает одно единственное значение.
- Замечательно! Вы верно поняли смысл доверительного интервала. Теперь самое время построить его!
- Мы строим двусторонний доверительный интервал!




---
## Ковариационная матрица

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: a458f30712



```

Попробуем посчитать ковариацию и корреляцию между временем ожидания извержения гейзера и его продолжительностью. Функция `cov()` позволяет получить ковариационную матрицу, где элемент матрицы $𝑎_{𝑖,𝑗}$ равен ковариации между *i* и *j* признаками. На главной диагонали расположены дисперсии значений признаков.

`@instructions`
- Постройте ковариацинную матрицу признаков.
- Посмотрите, как устроена матрица. Обратите внимание, что матрица симметричнаю
- Найдите ковариацию между признаками. Так как их всего два, подойдет любой элемент матрицы, не лежащий на главной диагонали. Обратиться к элементу матрицы можно с помощью квадратных скобок: `matrix[i, j]`.
- Найдите стандартное отклонение времени ожидания и продолжительности извержений.
- По формуле посчитайте корреляцию двух признаков. 
- Проверьте себя с помощью функции cor().

`@hint`
- В качестве аргумента функции ковариации используйте весь датасет. Из полученной матрицы извлеките ковариацию и посчитайте корни из дисперии (диагональные элементы).


`@sample_code`
```{r}
#Постройте ковариацинную матрицу признаков 
cov_matrix <- ___

#Посмотрите, как устроена матрица

#Найдите ковариацию между двумя признаками
covariance <- ___

#Найдите стандартное отклонение числа извержений
sd_erupt <- ___

#Найдите стандартное отклонение времени ожидания извержения гейзера
sd_wait <- ___

#Посчитатйте корреляцию времению ожидания и числа извержений
cor_1 <- ___

#В дальнейшем упрощайте себе жизнь, используя встроенную функцию корреляции
cor_2 <- ___

#Проверьте, что Вы получили одинаковые значения
```
`@solution`
```{r}
#Постройте ковариацинную матрицу признаков 
cov_matrix <-cov(faithful)

#Посмотрите, как устроена матрица 

#Найдите ковариацию между двумя признаками
covariance <- cov_matrix[2,1]

#Найдите стандартное отклонение числа извержений
sd_erupt <- sqrt(cov_matrix[1,1])

#Найдите стандартное отклонение времени ожидания извержения гейзера
sd_wait <- sqrt(cov_matrix[2,2])

#Посчитатйте корреляцию времению ожидания и числа извержений
cor_1 <- covariance/(sd_erupt*sd_wait)

#В дальнейшем упрощайте себе жизнь, используя встроенную функцию корреляции
cor_2 <- cor(faithful$eruptions, faithful$waiting)
```
`@sct`
```{r}
msg_1 = "Убедитесь, что Вы передали функции верный аргумент."
msg_2 = "Вам нужны элементы матрицы, не лежащие на диагонали."
msg_3 = "Извлечь квадртаный корень из числа можно при помощи функции sqrt()"
msg_4 = "Просмотрите формулу корреляции еще раз"
msg_5 = "В качестве аргументов передайте функции два вектора со значениями признаков"

test_object("cov_matrix", undefined_msg = msg_1, incorrect_msg = msg_1)
test_object("covariance", undefined_msg = msg_2, incorrect_msg = msg_2)
test_object("sd_erupt", undefined_msg = msg_3, incorrect_msg = msg_3)
test_object("sd_wait", undefined_msg = msg_3, incorrect_msg = msg_3)
test_object("cor_1", undefined_msg = msg_4, incorrect_msg = msg_4)
test_object("cor_2", undefined_msg = msg_5, incorrect_msg = msg_5)

test_error()
success_msg("Да Вы и правда молодец!")
```





---
## Доверительный интервал для математического ожидания

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: c7f8739f05



```

В этом задании мы построим 95 % доверительный интервал для среднего времени ожидания извержения гейзера. Так как мы имеем большое число наблюдений, можем воспользоваться универсальной формулой для математического ожидания. Большую часть работы Вы уже выполнили в предыдущих заданиях, поэтому здесь Вам останется только собрать все воедино. 

**Note:** В *R* существует несколько простых и удобных команд, которые облегчают работу с нормальным распределением:
- `rnorm(n, mean, sd)` - генерирует n нормально распределенных случайных величин.
- `qnorm(p, mean, sd)` - квантиль порядка p.
- `pnorm(q, mean, sd)` - вероятность, что случайная величина примет значение, меньше q.
- `dnorm(x, mead,sd)` - значение функции плотсности в точке x.

`@instructions`
- Найдите критическое значение статистики с помощью `qnorm()`. 
- Посчитатйе границы доверительного интервала. Сохраните полученные значения в переменные `low` и `high`.

`@hint`
- При уровне значимости $\alpha$ критические значения  соответствуют квантилям порядка
$\frac{\alpha}{2}$ и $\left(1-\frac{\alpha}{2}\right)$. В силу симметричности функции плотности стандартного нормального распределения эти значения равны по модулю.


`@sample_code`
```{r}
#Посчитайте среднее время ожидания извержения гейзера
mu_hat <- faithful$waiting

#Найдите критическое значение статистики
z_critic <- ___

#Посчитайте границы доверительного интервала. 
low <- ___

high <- ___
```
`@solution`
```{r}
#Посчитайте среднее время ожидания извержения гейзера
mu_hat <- faithful$waiting

#Найдите критическое значение статистики
z_critic <- qnorm(0.975)

#Посчитайте границы доверительного интервала
low <- mu_hat - z_critic*sd(faithful$waiting)/sqrt((faithful))

high <- mu_hat + z_critic*sd(faithful$waiting)/sqrt(nrow(faithful))
```
`@sct`
```{r}
msg_1 = "Не забывайте, что интервал двустороний."
msg_2 = "Кажется, Вы забыли формулу доверительного интервала."


test_object("z_critic", undefined_msg = msg_1, incorrect_msg = msg_1)
test_object("low", undefined_msg = msg_2, incorrect_msg = msg_2)
test_object("high", undefined_msg = msg_2, incorrect_msg = msg_2)
test_error()
success_msg("Вот это да! Вы осилили свой первый доверительный интервал. Двигаемся дальше!")
```





---
## Доверительный интервал для дисперсии

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 7b93b38b69



```

В заключительном задании воторой главы Вам предстоит построить 90% доверительный интервал для дисперсии времени извержения гейзера. Напомним, что статистика, которую Вы будете вычислять для построения доверительного интервала, имеет $\chi^2$ распределение с *n-1* степенями свободы. Для вычисления критических значений вам пригодится функция `qchisq(p, df)`, где `df`(*degree of freedom*) - число степеней свободы.

`@instructions`
- Найдите число степеней свободы. Сохраните значение в переменной `df`.
**Hint:** Длину вектора можно вычислить с помощью функции `length`.
- Посчитайте критические значения и сохраните их в переменный `crit_low` и `crit_high`. Они не симметричны, как это было в случае стандартного нормального распределения.
- Посчитайте верхнюю и нижнюю границы доверительного интервала.

`@hint`
- Количество наблюдений можно узнать с помощью функции `length()`. При уровне значимости $\alpha$ критические значения  соответствуют квантилям порядка
$\frac{\alpha}{2}$ и $\left(1-\frac{\alpha}{2}\right)$.


`@sample_code`
```{r}
#Посчитайте число степеней свободы
df <- ___

#Посчитайте критические значения
crit_low <- ___
crit_high <- ___


#Вычислите границы 90% доверительного интервала
high <- ___
low  <- ___
```
`@solution`
```{r}
#Посчитайте число степеней свободы
df <- length(df$eruptions) - 1

#Посчитайте критические значения
crit_low <- qchisq(0.95, df)
crit_high <- qchisq(0.05, df)


#Вычислите границы 90% доверительного интервала
high  <- (df-1)*var(df$eruptions)/crit_low
low   <- (df-1)*var(df$eruptions)/crit_high
```
`@sct`
```{r}
msg_1 = "Для хи-квадрат распределения число степеней свободы на единицу меньше числа наблюдений."
msg_2 = "Обратите внимание, что Ваш интервал лежит между двумя "хвостами"."


test_object("df", undefined_msg = msg_1, incorrect_msg = msg_1)
test_object("crit_low", undefined_msg = msg_2, incorrect_msg = msg_2)
test_object("crit_high", undefined_msg = msg_2, incorrect_msg = msg_2)

test_error()

success_msg("Ого! Это было последне задание в главе "Выборочные характеристики".)
```





---
## ECDF

```yaml
type: NormalExercise

xp: 100

key: ad84388c9f



```

В этой главе мы научимся вычислять основные выборочные характеристики при помощи *R*. Будем работать с датасетом `faithful`, который содержит данные о продолжительности `eruptions`  и времени ожидания `waiting` между извержениями гейзера Старый Служака. Данные уже доступны в рабочем пространстве.


Давайте посмотрим на распределение времени ожидания извержения гейзера. Для этого попробуем самостоятельно построить эмпирическую функцию распределения.

`@instructions`
- Прежде чем делать вычисления, полезно разобраться, что именно и в каких единицах измерения описывают Ваши переменные. Так как мы пользуемся  датасетом из пакета *R*, узнать это можно с помощью команды `?faithful`. Аналогичино, можно получить информацию о других объектах в *R*.
- Теперь можно перейти к анализу данных. Для построения красивых визуализации загрузите пакет `ggplot2`.
- Посчитатйте число наблюдений (количество строк в матрице `nrow()`).
- Постройте вариационный ряд значений времени ожидания. Вам пригодится функция `sort()`. Полученные значения сохраните в переменную `x_ecdf`.
- Сгенерируйте значения функции распределения и сохраните их в переменную `y_cdf`.
- С помощью функции `ggplot2()` постройте эмпирическую функцию распределения. Вам нужно указать оси `mapping = aes(x=  , y=   )`.

`@hint`
- Отсортируйте значения времени ожидания в порядке возрастасния. Затем сгенерируйте значения функции распределения. В данном случае они лежат в интервале $[\fraq{1}{n},1]$.


`@sample_code`
```{r}
#Познакомьтесь с описанием датасета

# Загрузите пакет для визуализации 
library(___)

#Найдите число наблюдений
n <- ___

#Постройте вариационный ряд значений времени ожидания
x_ecdf <- sort(___, decreasing = FALSE)

#Сгенерируйте значения выборочной функции распределения
y_ecdf <- seq(from =___, to= ___, by=___)/n

#Постройте выборочную функцию распределения
ggplot(mapping = aes(x=___ , y=___)) + geom_line() + labs(title = 'ECDF', x = 'Waiting time', y = 'F(x)'))
```
`@solution`
```{r}
#Загрузите пакет для визуализации данных
library(ggplot2)

#Найдите число наблюдений
n <- nrow(faithful)

#Постройте вариационный ряд значений времени ожидания
x_ecdf <- sort(faithful$waiting, decreasing = FALSE)

#Сгенерируйте значения выборочной функции распределения
y_ecdf <- seq(from = 1, to = n, by=1)/n

#Постройте выборочную функцию распределения
ggplot(mapping = aes(x_ecdf, y_ecdf)) + geom_line() + labs(title = 'ECDF', x = 'Waiting time', y = 'F(x)')
```
`@sct`
```{r}
msg_1 = "Число наблюдений равно числу строк в датасете"
msg_2 = "Нужно получить отсортированные значения времени ожидания"
msg_3 = "Какая область значений функции?"



test_object("n", undefined_msg = msg_1, incorrect_msg = msg_1)
test_object("x_ecdf", undefined_msg = msg_2, incorrect_msg = msg_2)
test_object("y_ecdf", undefined_msg = msg_3, incorrect_msg = msg_3)

test_error()
success_msg("Вы мастер выборочных функций распределения!")
```



