---
title       : Задание
description : Insert the chapter description here
---
## <<<Метод максимального правдоподобия>>>

```yaml
type: MultipleChoiceExercise
key: 70182e4c8d
lang: r
xp: 50
skills: 1
```


`@instructions`
Что необходимо знать для того, чтобы использовать для оценки параметра метод максимального правдоподобия?

`@hint`

`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```
