---
title: Теоретическая глава
description: This is a template chapter.
---

## Будет ли сегодня дождь?

```yaml
type: PureMultipleChoiceExercise
key: 4ff2227cf9
xp: 50
skills: 1
```
Известно, что каждый день дождь может либо идти, либо не идти. Каждое утро исследователь Геннадий включает специальный прибор, который выводит "1", если дождь в этот день будет, и "0", если дождя не будет. Выберите верное утверждение о работе прибора.

`@possible_answers`

- [Значение прибора до его включения – случайная величина.]
- Значение прибора до его включения – константа.
- Геннадий увидел, что прибор показывает "1". Значение прибора – случайная величина. 
- В прошлый четверг дождя не было, и утром прибор показал "0". Показанное значение прибора – случайная величина. 

`@hint`
Случаная величина может принимать некоторые значения с определенной вероятностью. Реализация случайной величины в данный момент времени – константа. 

`@feedback`
- Верно! До того, как прибор включён, Геннадий может только предположить, какое значение он покажет.
- Геннадий не знает, какую константу покажет прибор до включения.
- Геннадий знает, какое значение показал прибор, случайность отсутствует. 
- Геннадий знает, какое значение показал прибор, случайность отсутствует.




---
## В ожидании непогоды

```yaml
type: NormalExercise
key: 06813252d3
lang: r
xp: 100
skills: 1
```
Как вы поняли, значение прибора до включения – случайная величина, имеющая распределение Бернулли:

![](https://dl.dropboxusercontent.com/s/ocxkserxtg21gax/1_1.png?dl=0 =345x50)

В этом задании вам предстоит найти математическое ожидание значения прибора двумя способами:
* Непосредственно по математической формуле. Помните, что помимо прочего, R – очень удобный калькулятор;
* Используя команду `sum(x*p)`;

– а затем использовать полученное значение для нахождения дисперсии таким же образом.
Пусть Геннадий знает, что прибор выводит "1" с вероятностью 0.25.

`@instructions`

* Напоминаем, что формула для расчёта математического ожидания – это:
$$E(X) = \sum xp$$

  Эта формула уже внесена в код справа. Заполните её необходимыми числами.

* Теперь попробуем автоматизировать вычисления. Заполните вектор значений `x` случайной величины X, а также вектор вероятностей `p`. Рассчитайте математическое ожидание и сохраните его в переменную `mu`. Выведите `mu` и убедитесь, что результаты совпадают.

* Попробуем использовать полученное `mu` для нахождения диспресии. Напоминаем формулу:
$$Var(X) = E(X - E(X))^2$$

  Дополните код справа и сохраните полученное значение в переменную `var`. Выведите её.

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
# Заполните формулу необходимыми числами 
0 * (1 - ___) + 1 * ___

# Вектор значений x и вероятностей p
x <- c(___, ___)
p <- c(___, ___)

# Математическое ожидание
mu <- sum(___)
mu

# Дисперсия
var <- ___((x-___)^2*p)
___

```

`@solution`
```{r}
# Заполните формулу необходимыми числами 
0 * (1 - 0.25) + 1 * 0.25

# Вектор значений x и вероятностей p
x <- c(0, 1)
p <- c(0.75, 0.25)

# Математическое ожидание
mu <- sum(x*p)
mu

# Дисперсия
var <- sum((x-mu)^2*p)
var

```

`@sct`
```{r}

```
