---
  title: "Теория вероятностей и R"
  description: "This is a template chapter."
  v2: true

---
## Будет ли сегодня дождь?

```yaml
type: PureMultipleChoiceExercise

xp: 50
skills: 1
key: 4ff2227cf9



```

Известно, что каждый день дождь может либо идти, либо не идти. Каждое утро исследователь Геннадий включает специальный прибор, который выводит "1", если дождь в этот день будет, и "0", если дождя не будет. Выберите верное утверждение о работе прибора.


`@hint`
Случаная величина может принимать некоторые значения с определенной вероятностью. Реализация случайной величины в данный момент времени – константа.





`@possible_answers`
- [Значение прибора до его включения – случайная величина.]
- Значение прибора до его включения – константа.
- Геннадий увидел, что прибор показывает "1". Значение прибора – случайная величина. 
- В прошлый четверг дождя не было, и утром прибор показал "0". Показанное значение прибора – случайная величина.





---
## В ожидании непогоды

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 06813252d3



```

Как вы поняли, значение прибора до включения – случайная величина, имеющая распределение Бернулли:

![](https://dl.dropboxusercontent.com/s/ocxkserxtg21gax/1_1.png?dl=0 =345x50)

В этом задании вам предстоит найти математическое ожидание значения прибора двумя способами:
* Непосредственно по математической формуле. Помните, что помимо прочего, R – очень удобный калькулятор;
* Используя команду `sum(x*p)`;

– а затем использовать полученное значение для нахождения дисперсии таким же образом.
Пусть Геннадий знает, что прибор выводит "1" с вероятностью 0.25.

`@instructions`
* Напоминаем, что формула для расчёта математического ожидания – это:
 $E(X) = \sum x$

  Эта формула уже внесена в код справа. Заполните её необходимыми числами.

* Теперь попробуем автоматизировать вычисления. Заполните вектор значений `x` случайной величины X, а также вектор вероятностей `p`. Рассчитайте математическое ожидание и сохраните его в переменную `mu`. Выведите `mu` и убедитесь, что результаты совпадают.

* Попробуем использовать полученное `mu` для нахождения диспресии. Напоминаем формулу:
$$Var(X) = E(X - E(X))^2$$

  Дополните код справа и сохраните полученное значение в переменную `variance`. Выведите её.

`@hint`
Внесите в формулу вероятность "успеха". Для нахождения математического ожидания используйте вектор `x = c(0, 1)`. Для нахождения диспресии используйте команду `sum()`.


`@sample_code`
```{r}
# Заполните формулу необходимыми числами 
0 * (1 - ___) + 1 * ___

# Вектор значений x и вероятностей p
x <- c(___, ___)
p <- c(___, ___)

# Математическое ожидание
mu <- sum(___)
mu

# Дисперсия
variance <- ___((x-___)^2*p)
___
```
`@solution`
```{r}
# Заполните формулу необходимыми числами 
0 * (1 - 0.25) + 1 * 0.25

# Вектор значений x и вероятностей p
x <- c(0, 1)
p <- c(0.75, 0.25)

# Математическое ожидание
mu <- sum(x*p)
mu

# Дисперсия
variance <- sum((x-mu)^2*p)
variance
```
`@sct`
```{r}
ex() %>% check_output("0.25", missing_msg = "Проверьте, правильно ли Вы заполнили формулу математического ожидания.")
ex() %>% check_object("x") %>% check_equal(incorrect_msg = "Проверьте правильность заполнения вектора `x`.")
ex() %>% check_object("p") %>% check_equal(incorrect_msg = "Проверьте правильность заполнения вектора `p`.")
ex() %>% check_object("mu") %>% check_equal(incorrect_msg = "Проверьте правильнось переменной `mu`.")
ex() %>% check_object("variance") %>% check_equal()
ex() %>% check_output("0.1875", missing_msg = "Проверьте, вывели ли Вы полученное значение дисперсии.")
success_msg("Здорово! Как Вы видите, R очень удобно использовать в качестве продвинутого калькулятора. Но он способен на гораздо большее!")
```





---
## Бутерброд с маслом

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: d9a2b6ac19



```

Отложив прибор в сторону, исследователь Геннадий проводит свой утренний эксперимент: три раза роняет бутерброд с маслом, после чего записывает в специальную тетрадь результаты: на какую сторону упал бутерброд. Геннадий знает, что вероятность «приземления» на каждую сторону равна 0.5, и что каждое падение происходит независимо от других. В этом задании вам предстоит проанализировать один из таких экспериментов.

Обозначим за $X$ число падений на сторону с маслом. Как вы уже догадались, каждое падение можно представить собой как испытание Бернулли, а значит $X$ ~ $Bin(3, 0.5)$. Закон распределения выглядит следующим образом:

![](https://dl.dropboxusercontent.com/s/srykcpnabflq16a/1_2.png?dl=0)

`@instructions`
* Пришло время расширить ваши вычислительные возомжности. Рассчитайте вероятность $P(X=1)$, используя формулу, приведённую в таблице, и сохраните результат в переменную `p_1`. Биномиальный коэффициент $C_{n}^k$ можно найти при помощи команды `choose(n, k)`.

* Попробуем автоматизировать и это. Команда `dbinom(k, n, p)` вернёт вероятность того, что случайная величина, имеющая биномиальное распределение для `n` опытов Бернулли, в каждом из которых вероятность "успеха" составляет `p`, примет значение `k`. Используя эту команду, вычислите вероятность $P(X=2)$ и сохраните результат в `p_2`.

* Посмотрим, как связаны распределение Бернулли и Биномиальное распределение. Вычислите математическое ожидание одного падения. Сохраните его в переменную `mu_single`. Выведите `mu_single` и убедитесь, что она равна $p$.

* Вычислите математическое ожидание всех результатов эксперимента и сохраните полученное значение в переменную `mu_all`. Выведите `mu_all` и убедитесь, что она равна $np$.

* Вычислите дисперсию всех результатов эксперимента.Присвойте переменной `var_all` полученное значение. Выведите `var_all` и убедитесь, что она равна $np(1-p)$. 

* Рассчитайте стандартное отклонение всех результатов эксперимента и сохраните полученный ответ в переменную `sd`. Команда `sqrt()` может в этом пригодиться.



`@sample_code`
```{r}
# Вероятность P(X=1)
p_1 <- choose(3, ___)*___*0.5^___

# Вероятность P(X=2)
p_2 <- ___(___, 3, ___)

# Математическое ожидание одного падения
x_single <- c(___, ___)
p_single <- c(___, ___)
mu_single <- sum(___ * ___)
___

# Математическое ожидание всех падений
x_all <- c(___, ___, ___, 3)
p_all <- c(___, p_1, p_2, ___)
mu_all <- ___(x_all * ___)
mu_all

# Дисперсия всех падений
var_all <- ___((___ - ___)^___ * ___)
var_all

# Стандартное отклонение
___ <- ___(var_all)
```
`@solution`
```{r}
# Вероятность P(X=1)
p_1 <- choose(3, 1)*0.5*(0.5^2)

# Вероятность P(X=2)
p_2 <- dbinom(2, 3, 0.5)

# Математическое ожидание одного падения
#x_single <- c(0, 1)
#p_single <- c(0.5, 0.5)
#sum(x_single * p_single)

# Математическое ожидание всех падений
x_all <- c(0, 1, 2, 3)
p_all <- c(0.125, p_1, p_2, 0.125)
mu_all <- sum(x_all*p_all)

# Дисперсия всех падений
sum((x_all - mu_all)^2 * p_all)


# Стандартное отклонение
#sqrt(var_all)
```
`@sct`
```{r}
#ex() %>% check_object("p_1") %>% check_equal()
#ex() %>% check_object("p_2") %>% check_equal()
#ex() %>% check_output("0.5", missing_msg = "Проверьте, вывели ли Вы значение `mu_single`.")
#ex() %>% check_object("mu_all") %>% check_equal()
#ex() %>% check_output("1.5", missing_msg = "Проверьте, вывели ли Вы значение `mu_all`.")
#ex() %>% check_object("var_all") %>% check_equal()
#ex() %>% check_output("0.75", missing_msg = "Проверьте, вывели ли Вы значение `var_all`.")
#ex() %>% check_object("sd") %>% check_equal()
success_msg("Превосходно! Как вы заметили, Биномиальное распределение естественно происходит из распределения Бернулли. В следующем задании Вы сможете увидеть это.")
```





---
## Наглядно продемонстрируем!

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 758425b265



```

Распределение бутербродов, упавших маслом вниз, полученное вами в предыдущем задании, представлено здесь:

![](https://dl.dropboxusercontent.com/s/nbcziq9rafnoz82/1_3.png?dl=0)

`@instructions`
* Представьте закон распределения в виде векторов `x` и `p`.
* Для завершения обучения мастерству автоматизации расчётов с помощью R, найдите второй начальный и первый центральный моменты и сохраните полученные значения в переменные `second_raw` и `first_central` соответственно. Напоминаем, что
$$E(X^k)$$
$$E(X-E(X))^k$$
  \- k-ые начальный и центральный моменты.
* Попробуем построить функцию нашего распределения. Для этого будем использовать пакет `ggplot2`, который позволяет быстро строить красивые графики. Мы уже создали функцию, которая построит наше распределение. Сейчас вам не нужно понимать, как именно она работает (об этом вы узнаете в других курсах). Заполните её элементы `size` и `prob` числом испытаний эксперимента и вероятностью успеха одного испытания соответственно.

* Модифицируйте параметр `size`, чтобы построить функцию распределения Бернулли.


`@pre_exercise_code`
```{r}
library(ggplot2)
```
`@sample_code`
```{r}
# Закон распределения
x <- ___
p <- ___

# Второй начальный и первый центральный моменты
second_raw <- ___
mu <- ___
first_central <- ___

# Функция биномиального распределения
ggplot(data.frame(x = c(-5, 5)), aes(x)) + 
stat_function(fun = pbinom, args = list(size = ___, prob = ___), geom = 'step') + 
ylab("F(x)") +
ggtitle("Binomial Distribution") +
theme(plot.title = element_text(hjust = 0.5))

# Функция распределения Бернулли
ggplot(data.frame(x = c(-5, 5)), aes(x)) + 
stat_function(fun = pbinom, args = list(size = ___, prob = 0.5), geom = 'step') +
ylab("F(x)") +
ggtitle("Bernoulli Distribution") +
theme(plot.title = element_text(hjust = 0.5))
```
`@solution`
```{r}
# Закон распределения
x <- c(0, 1, 2, 3)
p <- c(0.125, 0.375, 0.375, 0.125)

# Второй начальный и первый центральный моменты
second_raw <- sum(x^2*p)
mu <- sum(x*p)
first_central <- sum((x-mu)*p)

# Функция биномиального распределения
ggplot(data.frame(x = c(-5, 5)), aes(x)) + 
stat_function(fun = pbinom, args = list(size = 3, prob = 0.5), geom = 'step') + 
ylab("F(x)") +
ggtitle("Binomial Distribution") +
theme(plot.title = element_text(hjust = 0.5))

# Функция распределения Бернулли
ggplot(data.frame(x = c(-5, 5)), aes(x)) + 
stat_function(fun = pbinom, args = list(size = 1, prob = 0.5), geom = 'step') +
ylab("F(x)") +
ggtitle("Bernoulli Distribution") +
theme(plot.title = element_text(hjust = 0.5))
```
`@sct`
```{r}
ex() %>% check_object("x") %>% check_equal()
ex() %>% check_object("p") %>% check_equal()
ex() %>% check_object("sencond_raw") %>% check_equal()
ex() %>% check_object("mu") %>% check_equal()
ex() %>% check_object("first_central") %>% check_equal()
ex() %>% {
  check_function(., "ggplot") %>% check_arg("data") %>% check_equal()
  check_function(., "aes") %>% {
    check_arg(., "x") %>% check_equal(eval = FALSE) 
    check_arg(., "y") %>% check_equal(eval = FALSE)
  }
  check_function(., "stat_function")
}
ex() %>% {
  check_function(., "ggplot") %>% check_arg("data") %>% check_equal()
  check_function(., "aes") %>% {
    check_arg(., "x") %>% check_equal(eval = FALSE) 
    check_arg(., "y") %>% check_equal(eval = FALSE)
  }
  check_function(., "stat_function")
}
success_msg("Чудесно! Посмотрите на графики: распределение Бернулли – это одна "ступенька" Биномиального распределения! А как будет выглядеть график при очень большом числе "ступенек"?)
```






---
## Функции плотности нормального распределения

```yaml
type: PureMultipleChoiceExercise

xp: 50

key: 7f571b03f8



```

Перед вами графики функций плотности нормального распределения с различными значениями математического ожидания и дисперсии. Подберите для графиков 1 (красный), 2 (зеленый) и 3 (синий) верные значения характеристик.
![](https://dl.dropboxusercontent.com/s/p26cqr3gzacl5q8/2_1.png?dl=0)


`@hint`
Обратите внимание на значения, в которых функции плотности достигают максимума, а также на то, все ли значения дисперсии допустимы.





`@possible_answers`
- 1) E(X) = 1, Var(X) = 0.2; 2)  E(X) = 1, Var(X) = 1; 3)  E(X) = 3, Var(X) = 5.
- [1)  E(X) = 0, Var(X) = 0.2; 2)  E(X) = 0, Var(X) = 5; 3)  E(X) = 2, Var(X) = 0.5.]
- 1)  E(X) = 0, Var(X) = 5; 2)  E(X) = 0, Var(X) = 0.2; 3)  E(X) = 2, Var(X) = 0.5.
- 1)  E(X) = 1, Var(X) = -0.2; 2)  E(X) = 1, Var(X) = -5; 3)  E(X) = 2, Var(X) = -1.

`@feedbacks`
- Попробуйте сравнить разброс значений случайной величины, соответствующий 2 и 3 графикам функций плотности, а также обратите внимание на значение математических ожиданий.
- Хорошая работа! Вы научились определять значения основных характеристик по виду функции плотности!
- Попробуйте сравнить разброс значений случайной величины, соответствующий 1 и 2 графикам функций плотности.
- Обратите внимание на допустимые значения дисперсии.


