---
  title: "Задание"
  description: "Insert the chapter description here"
  v2: true

---
## Метод максимального правдоподобия

```yaml
type: PureMultipleChoiceExercise

xp: 50
skills: 1
key: 0ef3cf6ecf



```

Что необходимо знать для того, чтобы использовать для оценки параметра метод максимального правдоподобия (ММП)?


`@hint`
Функция плотности может быть получена из функции распределения случайных величин.





`@possible_answers`
- Выборочное среднее; 
- [Функцию распределения случайных величин;] 
- Оценку параметра методом моментов; 
- Стандартное отклонение.

`@feedbacks`
- Попробуйте еще раз.
- Замечательно! Действительно, для применения ММП необходимо максимизировать функцию правдоподобия, для нахождения которой нужно знать, как случайные величины распределены.
- Попробуйте еще раз.
- Попробуйте еще раз.



---
## Вычисление квантилей

```yaml
type: MultipleChoiceExercise
key: 6c5bb1e1b1
lang: r
xp: 50
skills: 1
```
Определите, квантиль какого уровня больше для нормального распределения: 0.05 или 0.1? Почему? Нормальное распределение сгенерировано и доступно в рабочем пространстве под названием normal.


`@instructions`
- Квантиль уровня 0.05 больше, чем квантиль уровня 0.1; 
- [Квантиль уровня 0.1 больше, чем квантиль уровня 0.05;] 
- Квантили уровней 0.05 и 0.1 равны.

`@hint`
Чтобы сравнить квантили уровней p1 и p2, используйте следующее выражение: quantile(normal, p1) > quantile(normal, p2) 


`@pre_exercise_code`
normal <- rnorm(1000000, 0, 1)


`@sct`
- Попробуйте еще раз. Для вычисления квантилей уровня p, используйте функцию quantile(normal, p).
- Превосходно! Действительно, квантиль уровня 0.1 больше для нормального распределения, потому что мы с большей вероятностью попадем левее какого-либо числа (площадь под графиком функции плотности будет больше), если число, квантиль p-ого порядка, будет правее, то есть больше.
- Попробуйте еще раз. Для вычисления квантилей уровня p, используйте функцию quantile(normal, p).

