---
  title: "Выборочные характеристики"
  description: "This is a template chapter."
  v2: true

---
## Задание 2.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 2
key: dbe164cc29



```

В этой главе мы научимся вычислять основные выборочные характеристики при помощи *R*. Будем работать с датасетом `faithful`, который содержит данные о продолжительности `eruptions`  и времени ожидания `waiting` между извержениями гейзера Старый Служака. Данные уже доступны в рабочем пространстве. 
Давайте посмотрим на распределение времени ожидания извержения гейзера. Для этого опробуем самостоятельно построить эмпирическую функцию распределения.


**Note:**  Для построения Вам необходимо вспомнить, что такое вариационный ряд
$$ X_(n-1) \le X_(n-1) $$ ##дописать

`@instructions`
- Прежде чем делать вычисления, полезно разобраться, что именно и в каких единицах измерения описывают Ваши переменные. Узнать это можно с помощью команды `?faithful`. Аналогичино, можно получить информацию о других объъктах в *R*.
- Теперь можно перейти к анализу данных. Для построения красивых визуализации загрузите пакет `ggplot2`.
- Посчитатйте число наблюдений (количество строк в матрице `nrow()`).
- Постройте вариационный ряд значений времени ожидания. Вам пригодится функция `sort()`. Полученные значения сохраните в переменную `x_ecdf`.
- Сгенерируйте значения функции распределения и сохраните их в переменную `y_cdf`.
- С помощью функции `ggplot2()` постройте эмпирическую функцию распределения. Вам нужно указать оси `mapping = aes(x=  , y=   )`. Тип графика  - `geom_line()`.

`@hint`
- Here is the hint for this setup problem. It should get students 50% of the way to the correct answer.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
#library(datasets)
#datasets('faithful')
#data <- faithful
```
`@sample_code`
```{r}
#Познакомьтесь с описанием датасета

# Загрузите пакет для визуализации 
library(___)

#Найдите число наблюдений
n <- ___

#Постройте вариационный ряд значений времени ожидания
x_ecdf <- sort(___, decreasing = FALSE)

#Сгенерируйте значения выборочной функции распределения
y_ecdf <- seq(from =__ , to =__, by=___ )/n

#Постройте выборочную функцию распределения
ggplot(mapping = aes(x=__ ,y=___  )) + __ + labs(title = 'Эмпирическая функция распределения', x = 'Время ожидания', y = 'F(x)')+theme(plot.title = element_text(hjust = 0.5))
```
`@solution`
```{r}
#Загрузите пакет для визуализации данных
library(ggplot2)

#Найдите число наблюдений
n <- nrow(faithful)

#Постройте вариационный ряд значений времени ожидания
x_ecdf <- sort(faithful$waiting, decreasing = FALSE)

#Сгенерируйте значения выборочной функции распределения
y_ecdf <- seq(from = 1, to = n, by=1)/n

#Постройте выборочную функцию распределения
ggplot(mapping = aes(x_ecdf, y_ecdf)) + geom_line() + labs(title = 'Эмпирическая функция распределения', x = 'Время ожидания', y = 'F(x)')+theme(plot.title = element_text(hjust = 0.5))
```
`@sct`
```{r}
msg_1 = "В качества аргумента `nrow()` передайте название датасета"
msg_2 = "Обратиться к значениям переменной по названию можно с помощью `faithful$waiting`"
msg_3 = "Функция распределения может принимать значения от 0 до 1."
msg_4 = "Похоже, Вы неверно указали оси"
msg_5 = "Убедитесь, что вы верно указали тип графика `geom_line()`"


test_object("n", undefined_msg = msg_1, incorrect_msg = msg_1)
test_object("x_ecdf", undefined_msg = msg_2, incorrect_msg = msg_2)
test_object("y_ecdf",  undefined_msg = msg_3, incorrect_msg = msg_3)
test_output_contains(ggplot(mapping = aes(x_ecdf, y_ecdf)), undefined_msg = msg_3, incorrect_msg = msg_3)
test_functiont("geom_line()",  undefined_msg = msg_5, incorrect_msg = msg_5)
test_error()
success_msg("Great!")
```






---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 6bf93ce506
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## <<<Интерпретация доверительных интервалов>>>

```yaml
type: PureMultipleChoiceExercise
key: e692d3bd0c
xp: 50
skills: 1
```

Выберите верное суждение о 99% доверительном интервале
`@possible_answers`
- Построив 99 % доверительный интервал, мы получаем интервал, в котором лежит 99 % наших наблюдений.
- С вероятностью 99% неизвестный параметр  принимает значения из полученного доверительного интервала.
- [Проведя эксперимент большое число раз, мы ожидаем, что в 1% экспериментов доверительный интервал не будет накрывать истинное значение параметра.]
- Вероятность, что неизвестный параметр примет значение больше верхней границы доверительного интервала, равна 1%.


`@hint`


`@feedback`
"It is cool"
---
## Задание 2.2

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: a458f30712



```



`@instructions`
-Давайте нариусем пуассона
-ддда

`@hint`









